<a href="https://colab.research.google.com/github/shreya024/Spaceonova-Applications-of-Programming/blob/main/Day2/SV_602.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected without errors. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:

```python
NAME = "SHREYA GHOSH"
```

---

# Assignment - 2

## Q1
Given a dictionary *d*, write a code to get the value `I love Space!` and store it in **ans1**

<!-- grade_id: cell-f4d97d0f546b95c5, points: 0, grade: false, solution: true -->
```python
d = {1:{2:3},2:{1:['a','b',[1,2,3]],'3':["No","I love Space!","X"]},'5':'2'}
ans1=d[2]["3"][1]


print(ans1)
```

Output:
```
I love Space!
```

<!-- grade_id: cell-66b047b55925a728, points: 1, grade: true, solution: false -->
```python

```

## Q2
Given a list *l*, write a code to remove all duplicates and store the result as a list in a sorted order in **ans2**

<!-- grade_id: cell-7bcb12625616ddf7, points: 0, grade: false, solution: true -->
```python
l = [1,2,3,23,23,3,13,3,13,13,13,13,13,13,23,23,121,21,31,32,4,2,1]
s=set(l)
list_unique=list(s)
list_unique.sort()
ans2 = list_unique

print(ans2)
```

Output:
```
[1, 2, 3, 4, 13, 21, 23, 31, 32, 121]
```

<!-- grade_id: cell-ac35bcbdd65755c9, points: 1, grade: true, solution: false -->
```python

```

## Q3
Complete the function `leap_year()` which takes a 4 digit number as an argument `year` and **returns** whether the year is a leap year or not

<!-- grade_id: cell-c5576df343c11577, points: 0, grade: false, solution: true -->
```python
def leap_year(year):
    if(year%400==0 or year%4==0 and year%100!=0):
      return True
    else:
      return False
leap_year(1804)
```

Output:
```
True
```

<!-- grade_id: cell-535e31b449c845b8, points: 1, grade: true, solution: false -->
```python
#Tests (Must pass without error)
assert leap_year(2020)==True
assert leap_year(2021)==False
assert leap_year(2000)==True
assert leap_year(1900)==False
assert leap_year(1804)==True
```

## Q4

Complete the function `get_grade()` which takes a number as an argument `marks` and returns the grade according to the following criteria:

         Marks                                    Grade
         > 90                                       O
         > 80 and <= 90                             A
         >= 60 and <= 80                            B
         below 60                                   C

<!-- grade_id: cell-8b9e891026d35ce9, points: 0, grade: false, solution: true -->
```python
def get_grade(marks):
    
    if(marks>90):
      return 'O'
    elif(marks>80 and marks<=90):
      return 'A'
    elif(marks>=60 and marks<=80):
      return 'B'
    else:
      return 'C'
get_grade(99)
```

Output:
```
'O'
```

<!-- grade_id: cell-ed5363bca4e1e165, points: 1, grade: true, solution: false -->
```python
#Tests (Must pass without error)
assert get_grade(77)=='B'
assert get_grade(60)=='B'
assert get_grade(90)=='A'
assert get_grade(44)=='C'
assert get_grade(100)=='O'
assert get_grade(80)=='B'
```

## Q5

Complete the function `get_middle_digit()` which takes a number of odd length as an argument `num` and **returns** the number in the middle.

<!-- grade_id: cell-d953d773c76a4ded, points: 0, grade: false, solution: true -->
```python
def get_middle_digit(num):
    length=len(str(num))
    length=int(length/2)
    number=int(num/(10**length)%10)
    return number
get_middle_digit(777)
```

Output:
```
7
```

<!-- grade_id: cell-0eca0395959bc199, points: 1, grade: true, solution: false -->
```python
#Tests (Must pass without error)
assert get_middle_digit(777)==7
assert get_middle_digit(123454321)==5
assert get_middle_digit(1234321)==4
assert get_middle_digit(1234567890987654321)==0
assert get_middle_digit(1)==1
```

```python

```